Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Zezhi"
COLLABORATORS = ""
```

---

# Exercise 1 - Shell basics

Work through as much of the Software Carpentry [lesson on the Unix Shell](http://swcarpentry.github.io/shell-novice/) as you can.  Run through the Setup section just below, then open a terminal through Jupyter to run through the exercises.

After you have completed the first few sections of the tutorial, return to this notebook.  Execute all of the cells, and answer all of the questions.


## Setup - getting required files 

To get started, you'll need to have the required files in your directory.  Use `wget` to get them:

```python
!wget http://swcarpentry.github.io/shell-novice/data/shell-novice-data.zip
```

```python
!unzip -l shell-novice-data.zip
```

```python
!unzip shell-novice-data.zip
```

*Note*: you only need to do this once per session while using Jupyter on datanotebook.org.  You can open a terminal now and work through the steps, and return to this notebook a little later, and the files will be available either way.  That's because you're working in the same local directory.

**However**!  If you download this file, close your Jupyter session for the night, then come back tomorrow and open up a new Jupyter session on the server again, you'll need to get those sample files again.  Just execute the cells above to do it.

Okay, let's get on with the exercise!

## Navigating Files and Directories

As you work through this section of the tutorial, complete the steps here as well, using the `!` shell escape command.  Execute each cell as you go.

These steps aren't exactly the same as what's in the tutorial, where the file layout is a little different and where they're not using a notebook like we are.  That's okay.  Just consider this practice.

```python
!whoami
```

Output:
```
zezhi
```

```python
!pwd
```

Output:
```
/home/zezhi/Downloads
```

```python
!ls -F
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh  quotes.txt
data-shell/			 rstudio-0.99.903-amd64.deb
exercise-01.ipynb		 shell-novice-data.zip
quotations.txt			 shell-novice-data.zip.1
```

```python
!ls -F
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh  quotes.txt
data-shell/			 rstudio-0.99.903-amd64.deb
exercise-01.ipynb		 shell-novice-data.zip
quotations.txt			 shell-novice-data.zip.1
```

```python
!ls -F data-shell/
```

Output:
```
creatures/  Desktop/	north-pacific-gyre/  pizza.cfg	writing/
data/	    molecules/	notes.txt	     solar.pdf
```

```python
!ls -aF
```

Output:
```
./				 quotations.txt
../				 quotes.txt
Anaconda3-4.1.1-Linux-x86_64.sh  rstudio-0.99.903-amd64.deb
data-shell/			 shell-novice-data.zip
exercise-01.ipynb		 shell-novice-data.zip.1
.ipynb_checkpoints/
```

```python
!ls -af .
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh  shell-novice-data.zip.1
shell-novice-data.zip		 data-shell
..				 quotes.txt
.ipynb_checkpoints		 quotations.txt
.				 rstudio-0.99.903-amd64.deb
exercise-01.ipynb
```

What is the difference between the two previous cells, and what does the single dot mean?

<!-- grade_id: 156, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE


DIFFERENCE:

In ls -aF, ls -a means listing all the files including hidden files starting with '.'. ls -F appends symbols to filenames. These symbols show useful information about files. 

In ls -af ., ls -a is the same as the previous one, meaning listing all the files including hidden files starting with '.'.    ls -f means no sorting and no color.     '.' means in the current working directory. 



The single dot means in the current working directory.

```python
!ls -F ..
```

Output:
```
anaconda3/  Documents/	      Music/	 R/		 Videos/
Courses/    Downloads/	      Pictures/  Templates/
Desktop/    examples.desktop  Public/	 Untitled.ipynb
```

What do the double dots mean?

<!-- grade_id: 155, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE
It means moving up one dictionary level.

```python
!ls data-shell/north-pacific-gyre/2012-07-03/
```

Output:
```
goodiff		NENE01736A.txt	NENE01843A.txt	NENE01978B.txt	NENE02040Z.txt
goostats	NENE01751A.txt	NENE01843B.txt	NENE02018B.txt	NENE02043A.txt
NENE01729A.txt	NENE01751B.txt	NENE01971Z.txt	NENE02040A.txt	NENE02043B.txt
NENE01729B.txt	NENE01812A.txt	NENE01978A.txt	NENE02040B.txt
```

## Working with Files and Directories

The following cells come from the next section of the tutorial.

```python
!ls -F
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh  quotes.txt
data-shell/			 rstudio-0.99.903-amd64.deb
exercise-01.ipynb		 shell-novice-data.zip
quotations.txt			 shell-novice-data.zip.1
```

```python
!mkdir thesis
```

<!-- grade_id: 121, points: 1, grade: true, solution: false -->
```python
import os
assert "thesis" in os.listdir()
```

```python
!ls -F
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh  rstudio-0.99.903-amd64.deb
data-shell/			 shell-novice-data.zip
exercise-01.ipynb		 shell-novice-data.zip.1
quotations.txt			 thesis/
quotes.txt
```

You can't use the nano editor here in Jupyter, so we'll use the `touch` command to create an empty file instead.

```python
!touch thesis/draft.txt
```

<!-- grade_id: 123, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" in os.listdir("thesis")
```

```python
!ls -F thesis
```

Output:
```
draft.txt
```

Removing files and directories.

```python
!rm thesis/draft.txt
```

<!-- grade_id: 125, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" not in os.listdir("thesis")
```

```python
!rm -r thesis
```

```python
!rmdir thesis
```

Output:
```
rmdir: failed to remove 'thesis': No such file or directory
```

<!-- grade_id: 127, points: 1, grade: true, solution: false -->
```python
assert "thesis" not in os.listdir()
```

```python
!ls
```

Output:
```
Anaconda3-4.1.1-Linux-x86_64.sh  quotes.txt
data-shell			 rstudio-0.99.903-amd64.deb
exercise-01.ipynb		 shell-novice-data.zip
quotations.txt			 shell-novice-data.zip.1
```

Renaming and copying files.

```python
!touch draft.txt
```

<!-- grade_id: 129, points: 1, grade: true, solution: false -->
```python
assert "draft.txt" in os.listdir()
```

```python
!mv draft.txt quotes.txt
```

<!-- grade_id: 130, points: 1, grade: true, solution: false -->
```python
assert "quotes.txt" in os.listdir()
assert "draft.txt" not in os.listdir()
```

```python
!ls
```

```python
!cp quotes.txt quotations.txt
```

<!-- grade_id: 134, points: 1, grade: true, solution: false -->
```python
assert "quotes.txt" in os.listdir()
assert "quotations.txt" in os.listdir()
```

```python

```